# Devoir d'exemple — Python (nbgrader)

**Consignes.** Complétez chaque cellule marquée en remplaçant `raise NotImplementedError()` par votre solution.
N'ajoutez ni ne supprimez de cellules. Exécutez tout le notebook avant de soumettre (*Run → Run All Cells*).

Total : **26 points** (dont 4 en correction manuelle, 22 auto-notés).


## Exercice 1 — Somme d'une liste

Écrivez `somme_liste(lst)` qui renvoie la somme des éléments de `lst` (0 pour une liste vide).

<!-- grade_id: sol_somme, points: 0, grade: false, solution: true -->
```python
def somme_liste(lst):
    """Renvoie la somme des éléments de lst."""
    ### BEGIN SOLUTION
    return sum(lst)
    ### END SOLUTION

```

<!-- grade_id: test_somme, points: 2, grade: true, solution: false -->
```python
"""Vérifie Somme d'une liste (2 pts)."""
assert somme_liste([1, 2, 3]) == 6
assert somme_liste([]) == 0
### BEGIN HIDDEN TESTS
assert somme_liste([-2, 2, 5]) == 5
assert somme_liste([10]) == 10
### END HIDDEN TESTS

```

## Exercice 2 — Factorielle

Écrivez `factorielle(n)` (n ≥ 0) renvoyant n!. On rappelle 0! = 1.

<!-- grade_id: sol_fact, points: 0, grade: false, solution: true -->
```python
def factorielle(n):
    """Renvoie n! (n entier >= 0)."""
    ### BEGIN SOLUTION
    r = 1
    for k in range(2, n + 1):
        r *= k
    return r
    ### END SOLUTION

```

<!-- grade_id: test_fact, points: 3, grade: true, solution: false -->
```python
"""Vérifie Factorielle (3 pts)."""
assert factorielle(0) == 1
assert factorielle(5) == 120
### BEGIN HIDDEN TESTS
assert factorielle(1) == 1
assert factorielle(7) == 5040
### END HIDDEN TESTS

```

## Exercice 3 — Nombres premiers

Écrivez `est_premier(n)` qui renvoie `True` si `n` est premier, `False` sinon.

<!-- grade_id: sol_premier, points: 0, grade: false, solution: true -->
```python
def est_premier(n):
    """True si n est premier."""
    ### BEGIN SOLUTION
    if n < 2:
        return False
    k = 2
    while k * k <= n:
        if n % k == 0:
            return False
        k += 1
    return True
    ### END SOLUTION

```

<!-- grade_id: test_premier, points: 3, grade: true, solution: false -->
```python
"""Vérifie Nombres premiers (3 pts)."""
assert est_premier(2) is True
assert est_premier(15) is False
### BEGIN HIDDEN TESTS
assert est_premier(1) is False
assert est_premier(97) is True
assert est_premier(100) is False
### END HIDDEN TESTS

```

## Exercice 4 — Comptage de voyelles

Écrivez `compter_voyelles(s)` qui renvoie le nombre de voyelles (aeiouy) dans `s`, insensible à la casse.

<!-- grade_id: sol_voyelles, points: 0, grade: false, solution: true -->
```python
def compter_voyelles(s):
    """Nombre de voyelles dans s."""
    ### BEGIN SOLUTION
    return sum(1 for c in s.lower() if c in 'aeiouy')
    ### END SOLUTION

```

<!-- grade_id: test_voyelles, points: 2, grade: true, solution: false -->
```python
"""Vérifie Comptage de voyelles (2 pts)."""
assert compter_voyelles('Bonjour') == 3
assert compter_voyelles('') == 0
### BEGIN HIDDEN TESTS
assert compter_voyelles('XYZ') == 1
assert compter_voyelles('aeiouy') == 6
### END HIDDEN TESTS

```

## Exercice 5 — Inversion de chaîne

Écrivez `inverser(s)` qui renvoie la chaîne `s` à l'envers.

<!-- grade_id: sol_inverser, points: 0, grade: false, solution: true -->
```python
def inverser(s):
    """Renvoie s inversée."""
    ### BEGIN SOLUTION
    return s[::-1]
    ### END SOLUTION

```

<!-- grade_id: test_inverser, points: 2, grade: true, solution: false -->
```python
"""Vérifie Inversion de chaîne (2 pts)."""
assert inverser('abc') == 'cba'
assert inverser('') == ''
### BEGIN HIDDEN TESTS
assert inverser('radar') == 'radar'
assert inverser('Python') == 'nohtyP'
### END HIDDEN TESTS

```

## Exercice 6 — Fibonacci

Écrivez `fib(n)` renvoyant le n-ième terme de Fibonacci avec `fib(0)=0`, `fib(1)=1`.

<!-- grade_id: sol_fib, points: 0, grade: false, solution: true -->
```python
def fib(n):
    """n-ième terme de Fibonacci."""
    ### BEGIN SOLUTION
    a, b = 0, 1
    for _ in range(n):
        a, b = b, a + b
    return a
    ### END SOLUTION

```

<!-- grade_id: test_fib, points: 3, grade: true, solution: false -->
```python
"""Vérifie Fibonacci (3 pts)."""
assert fib(0) == 0
assert fib(1) == 1
assert fib(10) == 55
### BEGIN HIDDEN TESTS
assert fib(2) == 1
assert fib(15) == 610
### END HIDDEN TESTS

```

## Exercice 7 — Moyenne pondérée

Écrivez `moyenne_ponderee(notes, coefs)` (listes de même longueur non vide).

<!-- grade_id: sol_moyenne, points: 0, grade: false, solution: true -->
```python
def moyenne_ponderee(notes, coefs):
    """Moyenne des notes pondérée par coefs."""
    ### BEGIN SOLUTION
    return sum(n * c for n, c in zip(notes, coefs)) / sum(coefs)
    ### END SOLUTION

```

<!-- grade_id: test_moyenne, points: 3, grade: true, solution: false -->
```python
"""Vérifie Moyenne pondérée (3 pts)."""
assert moyenne_ponderee([10, 20], [1, 1]) == 15
### BEGIN HIDDEN TESTS
assert abs(moyenne_ponderee([8, 16], [1, 3]) - 14) < 1e-9
### END HIDDEN TESTS

```

## Exercice 8 — Classe `CompteBancaire`

Complétez la classe : `deposer(m)` ajoute `m` au solde ; `retirer(m)` retire `m` mais lève `ValueError` si le solde deviendrait négatif.

<!-- grade_id: sol_compte, points: 0, grade: false, solution: true -->
```python
class CompteBancaire:
    def __init__(self, solde=0):
        self.solde = solde

    def deposer(self, montant):
        ### BEGIN SOLUTION
        self.solde += montant
        ### END SOLUTION

    def retirer(self, montant):
        ### BEGIN SOLUTION
        if montant > self.solde:
            raise ValueError('solde insuffisant')
        self.solde -= montant
        ### END SOLUTION

```

<!-- grade_id: test_compte, points: 4, grade: true, solution: false -->
```python
"""Vérifie CompteBancaire (4 pts)."""
c = CompteBancaire(100)
c.deposer(50)
assert c.solde == 150
c.retirer(30)
assert c.solde == 120
### BEGIN HIDDEN TESTS
c = CompteBancaire()
assert c.solde == 0
try:
    c.retirer(999)
    assert False, 'devait lever ValueError'
except ValueError:
    pass
### END HIDDEN TESTS

```

## Exercice 9 — Question de cours (correction manuelle, 4 pts)

Expliquez la différence entre une **liste** et un **tuple** en Python, et donnez un cas d'usage de chacun.

<!-- grade_id: manual_reponse, points: 4, grade: true, solution: true -->
*Votre réponse ici (double-cliquez pour éditer) :*

YOUR ANSWER HERE